```python
# First import some necessary libraries
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
from sklearn.preprocessing import StandardScaler
from sklearn.decomposition import PCA
from sklearn.neighbors import KernelDensity
import matplotlib.pyplot as plt
from matplotlib.colors import LogNorm
from sklearn import mixture
up, down = True, False
np.set_printoptions(precision=3)
pd.set_option("expand_frame_repr", True)
pd.set_option("colheader_justify", "right")
pd.set_option("display.max_rows", None)
pd.set_option("display.width", 500)
pd.set_option("display.max_columns", 100)
pd.set_option("display.max_colwidth", 75)
pd.set_option("display.precision", 4)
```

<a id='Topic0'></a>

# Unsupervised Learning: Assignment 1

In this assignment, we're going to be covering dimensionality reduction with PCA. This week, we're going to use a dataset of spotify tracks.

    
Datasets often have has a huge number of different features for each data point. In the case of our spotify tracks dataset, we have features for each song such as popularity, duration, and "loudness." Applications that utilize data like this -- in this case recommendation systems in music streaming apps -- often need to reduce the number of features/dimensions to the data to make comparing each data point, or song in this case, computationally efficient. In this assignment, we are going to walk you through using PCA dimensionality reduction on a dataset of music tracks, and look into some useful tasks we can conduct with the resulting PCA object. 


<br>




```python
# Here, we read in the dataset and drop any tracks which might be missing values for a feature.

tracks_dataframe = pd.read_csv("./assets/tracks.csv")
tracks_dataframe = tracks_dataframe.dropna()

# Take a look at what's in the dataset!

tracks_dataframe.head()
```

```python
#As you can see, Each track has features such as length, loudness, and speechiness. 
#We don't need every feature, so we're going to make a dataframe with only the features we want. 
feature_names = ["popularity","duration_ms","danceability","energy","key","loudness","speechiness","acousticness","liveness","tempo"]

#Here, we select only these features to simplify our dataset. Take a look at the new data!
tracks_dataframe = tracks_dataframe[feature_names]
tracks_dataframe.head()
```

# Part 1: Creating a Principal Component Analysis Object

The first step in conducting a PCA analysis is to create a PCA object. In this part, we will walk you through the process and examine some attributes of the resulting object. 

# Task 1a.

The first step when preparing a dataset for PCA is to normalize the dataset. 

Conceptually, remember that what normalization does in each feature is:
1. Adjusts the values of a feature so that the feature has a mean of zero
2. Values of a feature will be adjusted to have a standard deivation of 1. 


Complete the function below to return a ndarray containing the normalized contents of `tracks_dataframe` features.

Hints:
 - Helpful Documentation: https://scikit-learn.org/stable/modules/generated/sklearn.preprocessing.StandardScaler.html
 - The two tools you need here are default StandardScaler and `.fit_transform()`
 

<!-- grade_id: cell-8a320d2f003a7663, points: 0, grade: true, solution: false -->
```python
task_id = "1a"
```

<!-- grade_id: cell-8beb5b3a5809337e, points: 0, grade: false, solution: true -->
```python


def task_1a_solution():
    normalized_data = None
    
    # YOUR CODE HERE
    raise NotImplementedError()
    return normalized_data


```

```python
# Use this cell to explore your solution.

task_1a_solution()
```

<!-- grade_id: cell-9aa45397d04fdbd3, points: 5, grade: true, solution: false -->
```python
print(f"Task {task_id} - AG tests")
stu_ans = task_1a_solution()

print(f"Task {task_id} - your answer:\n{stu_ans}")

assert isinstance(
    stu_ans, np.ndarray
), f"Task {task_id}: Your function should return an np.ndarray. "

assert (
    stu_ans.shape[1] == 10
), f"Task {task_id}: Your returned ndarray has an incorrect number of columns. "

assert np.isclose(
    np.mean(stu_ans, axis=0), 0.0
).all(), (
    f"Task {task_id}: The columns of your returned ndarray should have a zero mean. "
)

assert np.isclose(
    np.std(stu_ans, axis=0), 1.0
).all(), (
    f"Task {task_id}: The columns of your returned ndarray should have a unit std. "
)


del stu_ans
```

# Task 1b:

Now, create a PCA object using the normalized data you produced in task 1a. 

For the arguments, use n_components=6 and random_state = 42.

Helpful Documentation: https://scikit-learn.org/stable/modules/generated/sklearn.decomposition.PCA.html

<!-- grade_id: cell-b8bada8190f8cc9e, points: 0, grade: true, solution: false -->
```python
task_id = '1b'

```

<!-- grade_id: cell-0f464d733c3a8865, points: 0, grade: false, solution: true -->
```python

def task_1b_solution():
    normalized_data = task_1a_solution()
    
    # YOUR CODE HERE
    raise NotImplementedError()
    return pca
```

```python
# Use this cell to explore your solution.

task_1b_solution()
```

<!-- grade_id: cell-23f609514da02534, points: 5, grade: true, solution: false -->
```python
#here are some tests which may help troubleshoot your answer.


print(f"Task {task_id} - AG tests")
stu_ans = task_1b_solution()

print(f"Task {task_id} - your answer:\n{stu_ans}")


assert isinstance(
    stu_ans, PCA
), f"Task {task_id}: Your function should return a PCA object. "

assert (
    stu_ans.n_components == 6
), f"Task {task_id}: Your PCA object should be set to produce 6 components. "
assert stu_ans.components_.shape == (
    6,
    10,
), f"Task {task_id}: Your PCA fit on the grouped data should have produced 10 columns (one for each of the review_ features)."


del stu_ans
```

<br>
Great! We now have our PCA object. Now we're going to examine some neat things we can do with it. 

# Task 1c:

We know that each feature has a different weight in each principal component. 

Complete the function below to return the 3 features that have the highest absolute value weight for each of the
six principal components. This tells us which features "influence" that principal component the most.

- At a high level, this is what the function should do:
    - Get loadings of the principal components from the PCA object's .components_ attribute.
    - Iterate through each loading
    - For that loading, take the biggest 3 absolute value feature weights and add the names of the corresponding names to the list of results.

    
- Hints: 
   - np.argsort() and np.abs() are useful for obtaining the indices of the highest feature weights in a loading.
   - Recall that you can use index into the `feature_names` list we created above to reference feature names. For example, if the third loading in the  list has the highest absolute value weight, you can get the feature name with `feature_names[2]`.
   - You should add each list of 3 feature names to the end of "results". The list should be sorted from least to most weight (the last feature name in each list should have the highest weight) 
     


<!-- grade_id: cell-331651a3247c2d99, points: 0, grade: true, solution: false -->
```python
task_id = '1c'


```

<!-- grade_id: cell-3c0d8d23e188d9d6, points: 0, grade: false, solution: true -->
```python


def task_1c_solution():
    result = []
    pcaObject = task_1b_solution()
    # YOUR CODE HERE
    raise NotImplementedError()

    return (result)
```

```python
# Use this cell to explore your solution.

task_1c_solution()
```

<!-- grade_id: cell-989f932bb6565418, points: 20, grade: true, solution: false -->
```python

print(f"Task {task_id} - AG tests")
stu_ans = task_1c_solution()


print(f"Task {task_id} - your answer:\n{stu_ans}")


assert isinstance(
    stu_ans, list
), f"Task {task_id}: Your function should return a list. "

assert (
    len(stu_ans) == 6
), f"Task {task_id}: Your function should return an list of length 6. "


# Some hidden tests

del stu_ans
```

Take a look at the first 3 principal components. Think about how each component represents a certain aspect of a song. 

# Task 1d: Principal Components ratio:

We know that each Principal component accounts for a portion of the dataset's variance.

Complete the below function to return a value representing how much variance the first 5 principal components account for.

Hint: The explained_variance_ratio_ attribute of the PCA object is useful.

<!-- grade_id: cell-76fd79f2b370f9ec, points: 0, grade: true, solution: false -->
```python
task_id = '1d'

```

<!-- grade_id: cell-96b8a39e3e462b0b, points: 0, grade: false, solution: true -->
```python
def task_1d_solution():
    tracksPCAObject = task_1b_solution()
    cumulative_ratio = 0.0
    #YOUR CODE HERE
    # YOUR CODE HERE
    raise NotImplementedError()
    return cumulative_ratio
```

```python
# Use this cell to explore your solution.

task_1d_solution()
```

<!-- grade_id: cell-7952f861dc4c2f4a, points: 20, grade: true, solution: false -->
```python

print(f"Task {task_id} - AG tests")
stu_ans = task_1d_solution()

print(f"Task {task_id} - your answer:\n{stu_ans}")

assert isinstance(
    stu_ans, float
), f"Task {task_id}: Your function should return a float. "



del stu_ans
```

# Task 2: Analyzing a scree plot.

With the PCA object you created in task 1b, we can create a scree plot shown below, which can help us determine the number of principal components we need to account for an adequate amount of variance in the data.
We've also added a vertical line at eigenvalue 1.0 which will help you in this question.

```python

#Notice how we use normal explained_variance. You just used explained_variance_ratio_ in the previous part!

df = pd.DataFrame({'eigenvalue': np.sqrt(task_1b_solution().explained_variance_), 
                   'PC': ['01', '02', '03', '04', '05', '06']})

plt.plot(df['PC'], df['eigenvalue'], marker='o', color='c', linestyle='-')

plt.axhline(y=1.0, color='r', linestyle='--', label='Eigenvalue = 1.0')

plt.xlabel('PC')
plt.ylabel('Eigenvalue')
plt.title('Scree Plot')

plt.legend()

plt.show()

```

<!-- grade_id: cell-692072fcdb91ceba, points: 0, grade: true, solution: false -->
```python
task_id = '2'
```

Based on the Kaiser Heuristic, how many principal components should we use to get a good representation of the dataset?

Put your answer as the return value in the function below:


<!-- grade_id: cell-5b24c9c4a9932f43, points: 0, grade: false, solution: true -->
```python

def task_2_solution():
    solution = None

    # YOUR CODE HERE
    raise NotImplementedError()
    
    return solution
```

<!-- grade_id: cell-71b7f19ddf546893, points: 0, grade: true, solution: false -->
```python

print(f"Task {task_id} - AG tests")
stu_ans = task_2_solution()

assert isinstance(
    stu_ans, int
), f"Task {task_id}: Your function should return a int. "

```

# Part 3: Analyzing a biplot

Often, biplots are used to visualize PCA components. The code below uses the PCA object you created in part 1 to draw a biplot of the first 2 principal components. Run the cell below to create the graph. In this part, you will answer questions about the biplot.

```python


def biplot(score, coeff, maxdim, pcax, pcay, labels=None):
    """Routine to generate a high-quality biplot"""

    pca1 = pcax - 1
    pca2 = pcay - 1
    xs = score[:, pca1]
    ys = score[:, pca2]
    n = min(coeff.shape[0], maxdim)
    scalex = 2.0 / (xs.max() - xs.min())
    scaley = 2.0 / (ys.max() - ys.min())
    text_scale_factor = 1.5

    plt.figure(figsize=(10, 9))
    plt.scatter(xs * scalex, ys * scaley, s=1)
    xspecific = score[[20362, 103573, 76020], pca1]
    yspecific = score[[20362, 103573, 76020], pca2]
    ##xspecific = score[31:34, pca1]
    ##yspecific = score[31:34, pca2]
    plt.scatter(xspecific[0] * scalex, yspecific[0] * scaley, color='Yellow', marker='*', s=100)
    plt.scatter(xspecific[1] * scalex, yspecific[1] * scaley, color='Purple', marker='*', s=100)
    plt.scatter(xspecific[2] * scalex, yspecific[2] * scaley, color='Pink', marker='*', s=100)


    for i in range(n):
        plt.arrow(0, 0, coeff[i, pca1], coeff[i, pca2], color="r", alpha=0.5)
        if labels is None:
            plt.text(
                coeff[i, pca1] * text_scale_factor,
                coeff[i, pca2] * text_scale_factor,
                "Var" + str(i + 1),
                color="r",
                ha="center",
                va="center",
            )
        else:
            plt.text(
                coeff[i, pca1] * text_scale_factor,
                coeff[i, pca2],
                labels[i],
                color="r",
                ha="center",
                va="center",
            )
    plt.xlim(-1, 1)
    plt.ylim(-1, 1)
    plt.xlabel("PC{}".format(pcax))
    plt.ylabel("PC{}".format(pcay))
    plt.grid()
    plt.tight_layout()
    return


def make_biplot():
    X_spotify_normalized = task_1a_solution()## part_fa()
    pca = task_1b_solution()
    X_spotify_pca = pca.transform(X_spotify_normalized)
    
    biplot(
        X_spotify_pca, np.transpose(pca.components_[0:2, :]), 5, 1, 2, labels=feature_names
    )
    print(len(X_spotify_pca))

    return

make_biplot()
```

# Task 3a: 

For the first task involving the biplot, answer the following questions: 


- Which are the two most correlated variables?
- Which are the two least correlated variables?

Populate the provided lists with your solution. Put them in alphabetical order.

<!-- grade_id: cell-0935c94bbbd7e9f4, points: 0, grade: true, solution: false -->
```python
task_id = '3a'
```

<!-- grade_id: cell-5a50910b17be4131, points: 0, grade: false, solution: true -->
```python


def task_3a_solution():
    most_correlated = ["",""]
    least_correlated =  ["",""] 
    
    # YOUR CODE HERE
    raise NotImplementedError()

    return most_correlated, least_correlated


```

<!-- grade_id: cell-269bafbacd10bc9d, points: 5, grade: true, solution: false -->
```python

print(f"Task {task_id} - AG tests")
stu_ans = task_3a_solution()

assert isinstance(
    stu_ans, tuple
), f"Task {task_id}: Your function should return a tuple. "

assert len(stu_ans) == 2, f"Task {task_id}: Your function should return a tuple of length 2."

for i, item in enumerate(stu_ans):
    assert isinstance(item, list), f"Task {task_id}: Element {i} of the tuple should be a list."
    assert len(item) == 2, f"Task {task_id}: List {i} should contain exactly two elements."
    for j, val in enumerate(item):
        assert isinstance(val, str), f"Task {task_id}: Element {j} in list {i} should be a string."


del stu_ans

```

# Task 3b:
Of the three stars, one refers the popular pop song "Paris" by the chainsmokers, another refers to a popular Spanish pop rock song, and a third refers to an old orchestral song. Which refers to the old orchestral song? 


Put your answer as the string 'Pink', 'Yellow', or 'Purple' 

<!-- grade_id: cell-784598038c8b19b0, points: 0, grade: true, solution: false -->
```python
task_id = '3b'
```

<!-- grade_id: cell-95848026e119fe4a, points: 0, grade: false, solution: true -->
```python

def task_3b_solution():
    ans = None

    # YOUR CODE HERE
    raise NotImplementedError()
    return ans
```

```python
# Use this cell to explore your solution.

task_3b_solution()
```

<!-- grade_id: cell-c726271fb3df26b3, points: 5, grade: true, solution: false -->
```python

print(f"Task {task_id} - AG tests")
stu_ans = task_3b_solution()

assert isinstance(
    stu_ans, str
), f"Task {task_id}: Your function should return a str. "


del stu_ans
```

# Part 4: Kernel Density

Almost done! In the last part of this assignment, we will be using PCA components plotted into coordinate space in a kernel density object, which will allow us to glean interesting data about songs in the dataset. In this part, we will be computing the "probability" of the orchestral song existing in the dataset. 

The code below is a helper function which will give us the first 2 principal components in coordinate format.


```python
def get_PCA_coordinates():

    X_spotify_normalized = task_1a_solution()
    pca = task_1b_solution()
    X_spotify_pca = pca.transform(X_spotify_normalized)
    xs = X_spotify_pca[:, 0]
    ys = X_spotify_pca[:, 1]
    coordinates = np.column_stack((xs, ys))
    return coordinates
```

# Task 4a:
The first step is to create the kernel density object. Create one using the coordinates from our helper function. 

Use 'gaussian' for the kernel argument.

See the following documentation to see how to create the kernel density object: 
https://scikit-learn.org/dev/modules/generated/sklearn.neighbors.KernelDensity.html

Hint: Use `coordinates` as X.

<!-- grade_id: cell-a378358560e9eefd, points: 0, grade: true, solution: false -->
```python
task_id = '4a'


```

<!-- grade_id: cell-7c1d4f982ab7fda6, points: 0, grade: false, solution: true -->
```python
def task_4a_solution():
    coordinates = get_PCA_coordinates()
    kde = None

    #YOUR CODE HERE
    # YOUR CODE HERE
    raise NotImplementedError()

    return kde
```

```python
# Use this cell to explore your solution.

task_4a_solution()
```

<!-- grade_id: cell-55dd5aa2efbbac37, points: 5, grade: true, solution: false -->
```python

print(f"Task {task_id} - AG tests")
stu_ans = task_4a_solution()

assert isinstance(
    stu_ans, KernelDensity
), f"Task {task_id}: Your function should return a KernelDensity object. "


del stu_ans
```

# Task 4b

Now that we have the KDE object, we can calculate the probability of the existence of certain 
songs! Return the probability of the orchestral song in the above biplot.  

- Hints:
  - The orchestral song is the 76020th coordinate.
  - kde.score_samples(in the kde documentation) and `np.exp` are useful here.



<!-- grade_id: cell-149c6c66e299c572, points: 0, grade: true, solution: false -->
```python
task_id = '4b'

```

<!-- grade_id: cell-27064242e517e51b, points: 0, grade: false, solution: true -->
```python
def task_4b_solution():
    coordinates = get_PCA_coordinates()
    kde = task_4a_solution()

    probability = None
    #YOUR CODE HERE
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return probability
```

```python
# Use this cell to explore your solution.

task_4b_solution()
```

<!-- grade_id: cell-e61994c90f77a227, points: 10, grade: true, solution: false -->
```python
print(f"Task {task_id} - AG tests")
stu_ans = task_4b_solution()


print(f"Task {task_id} - your answer:\n{stu_ans}")

assert isinstance(
    stu_ans, float
), f"Task {task_id}: Your function should return a float "


del stu_ans
```

# Closing thoughts: 
Here is a Gaussian mixture model based off the first 2 Principal components. Where do you think this orchestral song might be in this plot: close to the center or farther out?


```python
def GMM():
    X_spotify_normalized = task_1a_solution()
    pca = task_1b_solution()
    X_spotify_pca = pca.transform(X_spotify_normalized)
    xs = X_spotify_pca[:, 0]
    ys = X_spotify_pca[:, 1]
    coordinates = np.column_stack((xs, ys))
    clf = mixture.GaussianMixture(n_components=2, covariance_type='full')
    clf.fit(coordinates)
    x = np.linspace(-10.0, 10.0)
    y = np.linspace(-10.0,10.0)
    X, Y = np.meshgrid(x, y)
    XX = np.array([X.ravel(), Y.ravel()]).T
    Z = -clf.score_samples(XX)
    Z = Z.reshape(X.shape)
    plt.figure(figsize=(8, 6))
    plt.title("Negative log-likelihood predicted by a GMM")

    CS = plt.contour(
        X, Y, Z, norm=LogNorm(vmin=1.0, vmax=1000.0), levels=np.logspace(0, 3, 10)
    )
    CB = plt.colorbar(CS, shrink=0.8)
    plt.scatter(coordinates[:, 0], coordinates[:, 1], 10)

    plt.tight_layout()

GMM()

```